# BLU02 - Exercises Notebook

```python
import hashlib # for grading

import os
import pandas as pd
import numpy as np
import datetime
from sklearn.model_selection import train_test_split
from sklearn.base import BaseEstimator, TransformerMixin
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import StandardScaler, RobustScaler
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_absolute_error
```

## 1 Read the Programs data (graded)

In this first exercise, we aim to create a single dataframe, combining all programs from all seasons.

With a caveat though: **we want to include seasons from the year 1950 onwards**.

<!-- grade_id: cell-e16abeb47fc4ea37, points: 0, grade: false, solution: true -->
```python
def read_season(folder_path, file_name):
    path = os.path.join(folder_path, file_name)
    return pd.read_csv(path)

def read_programs(folder_path):
    files = os.listdir(folder_path)
    # Create a list with the name of all files containing programs from
    # 1950 inclusive and onwards (just the filename, no complete path.)
    # files_from_1950: List[str] = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Create a list with the dataframes
    # seasons: List[pd.DataFrame] = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Use pd.concat to create a single dataframe.
    # programs: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Drop the column GUID.
    # programs = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove Duplicated lines.
    # YOUR CODE HERE
    raise NotImplementedError()
    # Set the index to be the column ProgramID
    # YOUR CODE HERE
    raise NotImplementedError()
    return programs

programs = read_programs(os.path.join('data', 'programs'))
```

<!-- grade_id: cell-2656708135a5a5e4, points: 2, grade: true, solution: false -->
```python
assert programs['Season'].min() == '1950-51'
assert programs['Season'].max() == '2016-17'
assert programs.index.name == 'ProgramID'
assert programs.shape == (7341, 2)
```

## 2 Read the Concerts data (graded)

Although we list all transformations step-by-step for the sake of clarity, we expect you to use method chaining.

<!-- grade_id: cell-e228c2d82463c6b4, points: 0, grade: false, solution: true -->
```python
def make_concerts(file_path): 
    # Read concerts data and drop the GUID column.
    # concerts: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Remember to_datetime? We need it here. We need to parse the columns Date and 
    # Time. Use pd.to_datetime(...).dt.date for the Date. 
    # then use the same logic to create the column Hour and Minute from Time column.
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove Duplicated lines.
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove all lines with empty Time column. Then also drop the Time column.
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return concerts

concerts = make_concerts(os.path.join('data','concerts.csv'))
```

<!-- grade_id: cell-3d21007e725ab889, points: 2, grade: true, solution: false -->
```python
assert concerts.shape == (20812, 8)
assert concerts.Date.min() == datetime.date(1842, 12, 7)
assert concerts.Date.max() == datetime.date(2017, 7, 7)
assert concerts.Date.max() == datetime.date(2017, 7, 7)
assert concerts['Hour'].mode().values[0] == 20
assert concerts['Minute'].mode().values[0] == 0
assert list(concerts.iloc[1537][['Hour', 'Minute']].values) == [20,30]
assert list(concerts.iloc[1201][['Hour', 'Minute']].values) == [20,15]
assert set(concerts.columns) == set([
    'ProgramID', 'ConcertID', 'EventType', 'Location', 'Venue', 'Date', 'Hour', 'Minute'
])
```

## 3 Combine Programs and Concerts data (graded)

Let's combine both dataframes into a single dataset, using an inner join.

<!-- grade_id: cell-a65f1464b4525a8b, points: 0, grade: false, solution: true -->
```python
# Remember that you want to join on the index of one of the dataframes.
# Join only the concerts with valid ProgramIDs
# nyp = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ac9aef3d5251e36c, points: 2, grade: true, solution: false -->
```python
assert nyp.shape == (12943, 10)
assert set(nyp.columns) == set([
    'ProgramID', 'ConcertID', 'EventType', 'Location', 'Venue',
    'Date', 'Hour', 'Minute', 'Orchestra', 'Season'
])
```

## 4 Read Works and Soloists data (graded)

We will read the two remaining pieces of data. 

Again, albeit the step-by-step description, we encourage you to use method chaining.

<!-- grade_id: cell-503e208490ff38e0, points: 0, grade: false, solution: true -->
```python
def read_works(file_path):
    # Read the works data.
    # works: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Remove the Intervals (attention to the values in the isInterval column).
    # works: pd.DataFrame = ...
    # YOUR CODE HERE
    raise NotImplementedError()
    # Select the columns ProgramID, WorkID, ComposerName, WorkTitle, Movement and ConductorName.
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove Duplicated lines.
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove all lines with empty Movement column.
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return works


def read_soloists(file_path):
    # Read the soloists data and drop GUID and MovementID Columns
    # YOUR CODE HERE
    raise NotImplementedError()
    ## Remove Duplicated lines.
    # YOUR CODE HERE
    raise NotImplementedError()
    return soloists


works = read_works('data/works.csv')
soloists = read_soloists('data/soloists.csv')
```

<!-- grade_id: cell-b8389314995f18ea, points: 2, grade: true, solution: false -->
```python
assert works.shape == (24320, 6)
assert set(works.columns) == set([
    'ProgramID', 'WorkID', 'ComposerName', 'WorkTitle', 'Movement', 'ConductorName'
])

assert soloists.shape == (50292, 5)
assert set(soloists.columns) == set([
   'ProgramID', 'WorkID', 'SoloistName', 'SoloistInstrument', 'SoloistRole'
])
```

## 5 Combine Works and Soloists (graded)

Like we did for Programs and Concerts, now we combine Works and Soloists.

<!-- grade_id: cell-08d9a086cc5646cf, points: 0, grade: false, solution: true -->
```python
# Combine both dataframes, again using an inner type of join. An work is identified by the pair
# ProgramId, WorkID
# works_and_soloists : pd.DataFrame = ....
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4d9f103dfffd311b, points: 2, grade: true, solution: false -->
```python
assert works_and_soloists.shape == (23578, 9)
assert set(works_and_soloists.columns) == set(
    [
        'ProgramID', 'WorkID', 'ComposerName', 'WorkTitle', 'Movement',
        'ConductorName', 'SoloistName', 'SoloistInstrument', 'SoloistRole'
    ]
)
```

## 6 Combine everything (graded)

The final goal here is to create a single dataframe.

<!-- grade_id: cell-ce1f05022e8cd63a, points: 0, grade: false, solution: true -->
```python
# Combine works_and_soloists and nyp into a single dataframe.
# You need to figure out the common column shared between the two dataframes
# nyp_merged = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ce29fa3aec1c244e, points: 2, grade: true, solution: false -->
```python
assert nyp_merged.shape == (27725, 18)
assert set(nyp_merged.columns) == set(
    [
       'ProgramID', 'ConcertID', 'EventType', 'Location', 'Venue', 'Date',
       'Hour', 'Minute', 'Orchestra', 'Season', 'WorkID', 'ComposerName', 'WorkTitle',
       'Movement', 'ConductorName', 'SoloistName', 'SoloistInstrument',
       'SoloistRole'
    ]
)
```